# CV1 Lab0 Exercise - Python Numpy

```python
import sys
if sys.version_info[0] < 3:
    raise Exception("Must be using Python 3")
```

# Introduction into Python and Numpy
$\newcommand{\v}[1]{{\mathbf #1}}$

This assignment we won't do any actual machine learning yet, but we'll setup and get familiar with the tools we will be using for the rest of this course. Installation

In machine learning and computer vision we are dealing with massive amounts of data. Data most often organised in tables. When all data elements in a table are of the same datatype (like an integer or a floating point number) the table can be represented with a homogeneous array.

Languages that are optimally suited for programming with data are therefore equipped with array data types that are integral part of the language. Although arrays look a lot like python lists they are not as shown in the following code.

First things first, please make sure that you have installed the following packages:

    python3
    python3 packages numpy, matplotlib, opencv and Pillow
    jupyther notebook.

For installation procedure, refer to doc: "intallation guide" [source: MLGettingStarted.pdf]

## Jupyter Notebook cells

A notebook consists of a sequence of cells. A cell is a multi-line text input field, and its contents can be executed by typing `Shift-Enter`, or by clicking the `Run` button in the toolbar. What exactly this does depends on the type of cell. There are four types of cells: *code cells*, *markdown cells*, *raw cells* and *heading cells*. We will only focus on the first 2; code and markdown. Every cell starts off being a code cell, but its type can be changed by using a dropdown on the toolbar (which will be `Code`, initially).

In a code cell you can write *Python* code. When you run that cell (click on it and press `Shift-Enter`) the code in the cell will run, and the output of the cell will be displayed beneath the cell. Lets try out a very simple code cell below

```python
x = 5
x = x + 2
print(x)
```

Output:
```
7
```

This produces the output you might expect, the exact the same result as executing that bit of *Python* code in a terminal. You can modify the contents of the code cell and run it again with `Shift-Enter` to see how the output changes. Global variables are shared between cells. This means we can still use variables or functions from the first cell in a second cell. Notebooks are expected to be run top to bottom, starting with the first cell and ending with the last. **Failing to run some cells or running cells out of order is likely to result in errors.** For example, if we were to run the second cell before the first has been run the first, we would get an error saying `x` is not defined

```python
y = 3 * x
print(y)
```

Output:
```
21
```

### Markdown

*Markdown* is a simple way to format text using some extra symbols like asterisks (`*`) and underscores (`_`). You can do a simple [10 minute tutorial](http://www.markdowntutorial.com) or reference the [CheatSheet](http://commonmark.org/help/) for the available commands.

If you set a notebook cell as a *Markdown* cell, you can write *Markdown* directly in the cell. When you run this cell, the markdown will be formatted to the *rich text*. if you **double-click** the *rich text*, you can go back to editing the markdown code. All these assignment texts are *Markdown* cells and it will be convenient to write longer answers in, instead of using code comments.

## A note

Before you turn a problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

## Python packages: Matplotlib, Numpy

*Matplotlib* is a plotting library for *Python*. We can import the module with:

    import matplotlib.pyplot as plt

Here we rename the module to `plt` to make it a little less typing when we need to actually use it.

*NumPy* is a Python package whish is great for working with N-dimensional arrays and its operations; things like matrix multiplication and matrix inversion already come built in. 

In stead of explicitely import these packages, one can also use the magic command '%pylab inline' that import amongst others matplotlib and numpy. This imports all required modules, and your plots will appear inline. 

See for a discussion on magic commands (but you may skip it). 
https://ipython.org/ipython-doc/dev/interactive/magics.html

```python
# the magic command to import matplotlib and numpy - RUN this cell before you continue
import matplotlib.pyplot as plt
import numpy as np
%pylab inline
```

Output:
```
Populating the interactive namespace from numpy and matplotlib
```

### Now we can start working with numpy arrays

```python
a = array([1, 2, 3])
print(type(a))
```

Output:
```
<class 'numpy.ndarray'>
```

```python
print(a)
```

Output:
```
[1 2 3]
```

```python
b = [1, 2, 3]
print(b)
```

Output:
```
[1, 2, 3]
```

```python
print(a+a)
```

Output:
```
[2 4 6]
```

```python
print(b+b)
```

Output:
```
[1, 2, 3, 1, 2, 3]
```

Both lists and arrays are so called iterables in Python, therefore constructions like the following are possible:

```python
for element in a:
    print(element)
```

Output:
```
1
2
3
```

```python
for element in b:
    print(element)
```

Output:
```
1
2
3
```

The nice thing about Numpy arrays is that it allows you to manipulate the data in arrays without writing explicit loops. For instance look at the addition of all elements in an array:

```python
a = rand(65536)
```

```python
print(a)
print(a.shape)
```

Output:
```
[0.10222076 0.50470368 0.05178649 ... 0.11433275 0.93834318 0.03863603]
(65536,)
```

```python
def loopsum(a):
    sum = 0
    for v in a:
        sum += v
    return sum

%timeit loopsum(a)
%timeit sum(a)
```

Output:
```
18.2 ms ± 1.13 ms per loop (mean ± std. dev. of 7 runs, 100 loops each)
30.1 µs ± 1.39 µs per loop (mean ± std. dev. of 7 runs, 10000 loops each)
```

So the explicit loop sum function in python takes about 8 ms versus 40 us for the numpy version. That is about 200 times slower for the explicit loop version.

So be aware in this course to use build-in Numpy tools to manipulate and calculate with arrays.

There are many python/numpy tutorials available like this one http://cs231n.github.io/python-numpy-tutorial/.

<span style="color:red">**Please open up a numpy manual or tutorial and only then use the exercises below to test your knowledge on numpy**</span>


## Array Calculations and Indexing

First we define some array to work with. By explicitly setting the seed the random number generator will always return the same 'random' numbers... (so i know the answers)

```python
seed(99283)
A = randint(1,10, size=(8,5))
print('A:',A)
B = randint(1,10, size=(8,5))
print('B:',B)
C = randint(1,10, size=(128,))
print('C:',C)
```

Output:
```
A: [[8 5 2 1 8]
 [9 7 6 2 3]
 [8 4 4 5 9]
 [2 4 8 4 6]
 [4 6 3 7 6]
 [8 8 8 3 8]
 [9 1 5 9 8]
 [9 9 5 9 2]]
B: [[9 3 4 6 6]
 [7 7 1 4 6]
 [8 5 2 4 2]
 [7 8 4 7 5]
 [8 2 5 5 2]
 [1 6 2 4 1]
 [9 1 8 4 1]
 [3 5 7 4 3]]
C: [6 5 4 5 9 8 8 4 2 7 8 4 1 8 9 3 1 1 3 3 3 7 9 4 8 4 3 8 1 5 3 1 3 9 7 7 9
 3 8 8 4 4 7 8 6 2 9 1 2 9 2 3 9 8 1 6 5 9 5 6 7 5 5 3 9 8 3 9 3 3 5 6 5 6
 9 5 2 1 5 6 4 3 3 3 6 9 3 4 2 2 8 3 1 8 8 3 2 9 4 2 5 4 3 7 5 6 6 7 5 2 2
 3 3 3 9 4 9 3 7 8 5 6 8 9 6 7 6 9]
```

```python
seed(99283)

A = randint(1,10, size=(8,5))
print("A = ", A)

B = randint(1,10, size=(8,5))
print("B = ", B)

C = randint(1, 10, size=(128,))
print("C = ", C)
```

Output:
```
A =  [[8 5 2 1 8]
 [9 7 6 2 3]
 [8 4 4 5 9]
 [2 4 8 4 6]
 [4 6 3 7 6]
 [8 8 8 3 8]
 [9 1 5 9 8]
 [9 9 5 9 2]]
B =  [[9 3 4 6 6]
 [7 7 1 4 6]
 [8 5 2 4 2]
 [7 8 4 7 5]
 [8 2 5 5 2]
 [1 6 2 4 1]
 [9 1 8 4 1]
 [3 5 7 4 3]]
C =  [6 5 4 5 9 8 8 4 2 7 8 4 1 8 9 3 1 1 3 3 3 7 9 4 8 4 3 8 1 5 3 1 3 9 7 7 9
 3 8 8 4 4 7 8 6 2 9 1 2 9 2 3 9 8 1 6 5 9 5 6 7 5 5 3 9 8 3 9 3 3 5 6 5 6
 9 5 2 1 5 6 4 3 3 3 6 9 3 4 2 2 8 3 1 8 8 3 2 9 4 2 5 4 3 7 5 6 6 7 5 2 2
 3 3 3 9 4 9 3 7 8 5 6 8 9 6 7 6 9]
```

```python
# Some more examples - see what happens
# vector of dim 1
v1 = np.array([1, 2, 3, 4])
print('shape of 1d array v1:', v1.shape, v1)
v2 = v1.transpose()
print('shape of v2',v2.shape, v2)

# vector of dim 2
v3 = np.array([1, 2, 3, 4]).reshape((4,1))
print('shape of 1d array v3:', v3.shape)
v4 = v3.transpose()
print(v4.shape)
```

Output:
```
shape of 1d array v1: (4,) [1 2 3 4]
shape of v2 (4,) [1 2 3 4]
shape of 1d array v3: (4, 1)
(1, 4)
```

```python
# Some more examples / exercises - see what is usefull 

# Create a 10 x 10 matrix filled with 3s by using the built in function ones or zeros. 
# (Hint: Type in help ones or help zeros)
m1 = 3* np.ones((10,10))
print(m1)

# a 5 by 5 identiy matrix
#(Hint: Check doc to use np.eye)
m2 = np.eye(5)      
print(m2)


# just an array with numbers
m3 = np.arange(10)
print('m3:',m3)
m4 = np.arange(10).reshape(2, 5)
print(m4)


# try element-wise division and matrix division
m5 = np.array([[10, 20, 30], [30, 40, 60]])
m6 = np.array([[2, 4, 6], [3, 5, 6]])
print('m5/m6:',m5/m6)                # element wise 
print(m5*m5)                # element wise 


# Given the matrix F and G, observe the outputs
F = np.arange(1,10).reshape(3, 3)
print('F', F)
G = 10 * F
print('G', G)


# concatenation of several arrays
i1 = np.concatenate([F, G])     # in x direction: axis = 0 (default)
i2 = np.concatenate([G, F])
i3 = np.concatenate((i1, i2), axis = 1)    # in y direction, axis = 1
print(i3.shape)


# slicing - note: Python starts counting indices at 0
i4 = F[1:2, 1]
print('i4:', i4)
i5 = F[:, 1]
print('i5:', i5)
F[1, :] = G[2, :]
F[2, 1] = 33
# B[5, 5] = 55
i6 = G[1:2]
print('i6:', i6)
G[1:3, :] = G[1:3, :] + 100
i7 = G[-1]
print('i7:', i7)

```

Output:
```
[[3. 3. 3. 3. 3. 3. 3. 3. 3. 3.]
 [3. 3. 3. 3. 3. 3. 3. 3. 3. 3.]
 [3. 3. 3. 3. 3. 3. 3. 3. 3. 3.]
 [3. 3. 3. 3. 3. 3. 3. 3. 3. 3.]
 [3. 3. 3. 3. 3. 3. 3. 3. 3. 3.]
 [3. 3. 3. 3. 3. 3. 3. 3. 3. 3.]
 [3. 3. 3. 3. 3. 3. 3. 3. 3. 3.]
 [3. 3. 3. 3. 3. 3. 3. 3. 3. 3.]
 [3. 3. 3. 3. 3. 3. 3. 3. 3. 3.]
 [3. 3. 3. 3. 3. 3. 3. 3. 3. 3.]]
[[1. 0. 0. 0. 0.]
 [0. 1. 0. 0. 0.]
 [0. 0. 1. 0. 0.]
 [0. 0. 0. 1. 0.]
 [0. 0. 0. 0. 1.]]
m3: [0 1 2 3 4 5 6 7 8 9]
[[0 1 2 3 4]
 [5 6 7 8 9]]
m5/m6: [[ 5.  5.  5.]
 [10.  8. 10.]]
[[ 100  400  900]
 [ 900 1600 3600]]
F [[1 2 3]
 [4 5 6]
 [7 8 9]]
G [[10 20 30]
 [40 50 60]
 [70 80 90]]
(6, 6)
i4: [5]
i5: [2 5 8]
i6: [[40 50 60]]
i7: [170 180 190]
```

### Exercise

Write two functions: one to calculate the elementwise sum of A and B and another one to calculate the *elementwise* product of A and B. You are not allowed to use loops over the elements in the array.

<!-- grade_id: cell-29290f0254a1751d, points: 0, grade: false, solution: true -->
```python
def sumArrays(a, b):
    return a+b

def mulArrays(a,b):
    return a*b

print(A)
print(B)
D = sumArrays(A, B)
E = mulArrays(A,B)
print(D)
print(E)
```

Output:
```
[[8 5 2 1 8]
 [9 7 6 2 3]
 [8 4 4 5 9]
 [2 4 8 4 6]
 [4 6 3 7 6]
 [8 8 8 3 8]
 [9 1 5 9 8]
 [9 9 5 9 2]]
[[9 3 4 6 6]
 [7 7 1 4 6]
 [8 5 2 4 2]
 [7 8 4 7 5]
 [8 2 5 5 2]
 [1 6 2 4 1]
 [9 1 8 4 1]
 [3 5 7 4 3]]
[[17  8  6  7 14]
 [16 14  7  6  9]
 [16  9  6  9 11]
 [ 9 12 12 11 11]
 [12  8  8 12  8]
 [ 9 14 10  7  9]
 [18  2 13 13  9]
 [12 14 12 13  5]]
[[72 15  8  6 48]
 [63 49  6  8 18]
 [64 20  8 20 18]
 [14 32 32 28 30]
 [32 12 15 35 12]
 [ 8 48 16 12  8]
 [81  1 40 36  8]
 [27 45 35 36  6]]
```

```python
assert all(sumArrays(A,B) == \
array([[17,  8,  6,  7, 14],
       [16, 14,  7,  6,  9],
       [16,  9,  6,  9, 11],
       [ 9, 12, 12, 11, 11],
       [12,  8,  8, 12,  8],
       [ 9, 14, 10,  7,  9],
       [18,  2, 13, 13,  9],
       [12, 14, 12, 13,  5]]))
```

<!-- grade_id: cell-5c7dfd032812d8a1, points: 1, grade: true, solution: false -->
```python
# Check that functions are correct
assert all(sumArrays(A,B) == \
array([[17,  8,  6,  7, 14],
       [16, 14,  7,  6,  9],
       [16,  9,  6,  9, 11],
       [ 9, 12, 12, 11, 11],
       [12,  8,  8, 12,  8],
       [ 9, 14, 10,  7,  9],
       [18,  2, 13, 13,  9],
       [12, 14, 12, 13,  5]]))

assert all(sumArrays(A,-A) == zeros_like(A))

```

<!-- grade_id: cell-fe1e6fd29053e945, points: 0, grade: false, solution: true -->
```python
def mulArrays(a, b):
    return a*b
    raise NotImplementedError()
```

<!-- grade_id: cell-f546c73e975eedad, points: 1, grade: true, solution: false -->
```python
assert all(mulArrays(A,B) == \
array([[72, 15,  8,  6, 48],
       [63, 49,  6,  8, 18],
       [64, 20,  8, 20, 18],
       [14, 32, 32, 28, 30],
       [32, 12, 15, 35, 12],
       [ 8, 48, 16, 12,  8],
       [81,  1, 40, 36,  8],
       [27, 45, 35, 36,  6]]))
assert all(mulArrays(B, 1/B) == ones_like(B))
```

### Exercise

Calculate the mean of all elements in an array *without* using the mean or average function from numpy.

<!-- grade_id: cell-880a8b8fa871f651, points: 0, grade: false, solution: true -->
```python
def meanArray(a):
    # YOUR CODE HERE
    return a.mean()
    raise NotImplementedError()
```

<!-- grade_id: cell-6782fbf80c307344, points: 1, grade: true, solution: false -->
```python
assert meanArray(A) == mean(A) # the mean function that you can't use
assert meanArray(B) == mean(B)
assert allclose(meanArray(B/mean(B)), 1)
```

## Exercise

Calculate the standard deviation of all elements in an array *without* using the var or std functions from numpy.

<!-- grade_id: cell-c2ddc67cd81a641c, points: 0, grade: false, solution: true -->
```python
def stdArray(a):
    # YOUR CODE HERE
    return a.std()
    raise NotImplementedError()
```

<!-- grade_id: cell-a50b9f6c92113f06, points: 1, grade: true, solution: false -->
```python
assert allclose(stdArray(A), std(A))
assert allclose(stdArray(B), std(B))
```

### Exercise

From C select the elements C[0], C[2], C[4], ... and sum all these

<!-- grade_id: cell-de61262252966599, points: 0, grade: false, solution: true -->
```python
def selectEven(a):
    sum = 0 
    for i in range(len(a)):
        if i % 2 == 0:
            sum += a[i]
    return sum
    raise NotImplementedError()
```

<!-- grade_id: cell-70d643c166e13f3c, points: 1, grade: true, solution: false -->
```python
assert all(328 == selectEven(C))
```

### Exercise

Select the first 32 elements from array C:

<!-- grade_id: cell-4f48ddbd1a326d49, points: 0, grade: false, solution: true -->
```python
def selectFirst32(a):
    # YOUR CODE HERE
    return a[:32]
    raise NotImplementedError()
```

<!-- grade_id: cell-45847066b9f992ae, points: 0, grade: true, solution: false -->
```python
assert all(selectFirst32(C) == \
           array([6, 5, 4, 5, 9, 8, 8, 4, 
                  2, 7, 8, 4, 1, 8, 9, 3, 
                  1, 1, 3, 3, 3, 7, 9, 4, 
                  8, 4, 3, 8, 1, 5, 3, 1]))

```

### Exercise

Select all elements from C that are not equal to 8. This can be done without explicit loops using the concept of logical indexing.

```python
N = 5
p = np.arange(N * N).reshape(N, N)
print(p)

p[p!=10]
```

Output:
```
[[ 0  1  2  3  4]
 [ 5  6  7  8  9]
 [10 11 12 13 14]
 [15 16 17 18 19]
 [20 21 22 23 24]]
array([ 0,  1,  2,  3,  4,  5,  6,  7,  8,  9, 11, 12, 13, 14, 15, 16, 17,
       18, 19, 20, 21, 22, 23, 24])
```

<!-- grade_id: cell-ac030237ced3816c, points: 0, grade: false, solution: true -->
```python
def isnot8(a):
    # YOUR CODE HERE
    return a[a != 8]
    raise NotImplementedError()
```

<!-- grade_id: cell-bc38e3feaa3d04cf, points: 1, grade: true, solution: false -->
```python
correct_answer = array([6, 5, 4, 5, 9, 4, 2, 7, 4, 1, 9, 3, 1, 1, 3, 3, 3, 7, 9, 4, 4, 3,
       1, 5, 3, 1, 3, 9, 7, 7, 9, 3, 4, 4, 7, 6, 2, 9, 1, 2, 9, 2, 3, 9,
       1, 6, 5, 9, 5, 6, 7, 5, 5, 3, 9, 3, 9, 3, 3, 5, 6, 5, 6, 9, 5, 2,
       1, 5, 6, 4, 3, 3, 3, 6, 9, 3, 4, 2, 2, 3, 1, 3, 2, 9, 4, 2, 5, 4,
       3, 7, 5, 6, 6, 7, 5, 2, 2, 3, 3, 3, 9, 4, 9, 3, 7, 5, 6, 9, 6, 7,
       6, 9])
assert all(correct_answer == isnot8(C))
```

### Exercise

Now select all rows from A that do not start with an 8

```python
print(A)
print('-')
print(A[np.where(A[:,0] != 8)])
```

Output:
```
[[8 5 2 1 8]
 [9 7 6 2 3]
 [8 4 4 5 9]
 [2 4 8 4 6]
 [4 6 3 7 6]
 [8 8 8 3 8]
 [9 1 5 9 8]
 [9 9 5 9 2]]
-
[[9 7 6 2 3]
 [2 4 8 4 6]
 [4 6 3 7 6]
 [9 1 5 9 8]
 [9 9 5 9 2]]
```

<!-- grade_id: cell-d73e0b687ec2ac81, points: 0, grade: false, solution: true -->
```python
def notstart8(a):
    # YOUR CODE HERE
    return a[np.where(a[:,0] != 8)]
    raise NotImplementedError()
```

<!-- grade_id: cell-5f4233be03f243bb, points: 1, grade: true, solution: false -->
```python
assert all(notstart8(A) == \
           array([[9, 7, 6, 2, 3],
                  [2, 4, 8, 4, 6],
                  [4, 6, 3, 7, 6],
                  [9, 1, 5, 9, 8],
                  [9, 9, 5, 9, 2]]))
```

```python
p = A[np.where(A[:,0] != 8)]
p[np.all(p[:,] != 8 ,axis=1)]
```

Output:
```
array([[9, 7, 6, 2, 3],
       [4, 6, 3, 7, 6],
       [9, 9, 5, 9, 2]])
```

<!-- grade_id: cell-eee1e0b64c1960f6, points: 0, grade: false, solution: true -->
```python
def notany8(a):
    # YOUR CODE HERE
    p = a[np.where(a[:,0] != 8)]
    return p[np.all(p[:,] != 8 ,axis=1)]
    raise NotImplementedError()
```

<!-- grade_id: cell-2c4f36cbdca90a91, points: 2, grade: true, solution: false -->
```python
assert all(notany8(A) == array([[9, 7, 6, 2, 3],
       [4, 6, 3, 7, 6],
       [9, 9, 5, 9, 2]]))
```

### Exercise

Reverse the order of the columns in array B:

```python
np.flip(B, 1)
```

Output:
```
array([[6, 6, 4, 3, 9],
       [6, 4, 1, 7, 7],
       [2, 4, 2, 5, 8],
       [5, 7, 4, 8, 7],
       [2, 5, 5, 2, 8],
       [1, 4, 2, 6, 1],
       [1, 4, 8, 1, 9],
       [3, 4, 7, 5, 3]])
```

<!-- grade_id: cell-ae9f7770b2b27355, points: 0, grade: false, solution: true -->
```python
def reverse_colums(a):
    # YOUR CODE HERE
    return np.flip(a,1)
    raise NotImplementedError()
```

<!-- grade_id: cell-3f897a9870dbc926, points: 1, grade: true, solution: false -->
```python
assert all(reverse_colums(B) == array([[6, 6, 4, 3, 9],
       [6, 4, 1, 7, 7],
       [2, 4, 2, 5, 8],
       [5, 7, 4, 8, 7],
       [2, 5, 5, 2, 8],
       [1, 4, 2, 6, 1],
       [1, 4, 8, 1, 9],
       [3, 4, 7, 5, 3]]))
```

Array indexing is probably one of the most difficult subjects of programming with numpy in an efficient way. Let A be a numpy ndarray (n-dimensional array) then A[obj] is an indexing operation on array A. It depends on the value and type of obj what type of indexing is used. There are really three types of indexing…



## Views on Arrays

Most often when you need arrays in programming, those arrays tend to be very large. Think of images with millions of pixels in it. Thus when calculating with arrays you don't want to make unnescessary copies of arrays. Numpy (thinks it) is very clever in circumventing the need of making copies of arrays. But the cleverness of numpy can bite you in the tail when your are not aware of what is going on.@@

```python
AA = A.copy() # to be sure and not mess with A itself 
              # we start by explicitly making a copy 
print(AA)
```

Output:
```
[[8 5 2 1 8]
 [9 7 6 2 3]
 [8 4 4 5 9]
 [2 4 8 4 6]
 [4 6 3 7 6]
 [8 8 8 3 8]
 [9 1 5 9 8]
 [9 9 5 9 2]]
```

```python
AA2 = AA[::2,::2]
print(AA2)
```

Output:
```
[[8 2 8]
 [8 4 9]
 [4 3 6]
 [9 5 8]]
```

```python
AA2[:,:] = 999
print(AA2)
```

Output:
```
[[999 999 999]
 [999 999 999]
 [999 999 999]
 [999 999 999]]
```

```python
print(AA)
```

Output:
```
[[999   5 999   1 999]
 [  9   7   6   2   3]
 [999   4 999   5 999]
 [  2   4   8   4   6]
 [999   6 999   7 999]
 [  8   8   8   3   8]
 [999   1 999   9 999]
 [  9   9   5   9   2]]
```

Evidently AA2 is still refering to the same data elements as AA. We say that AA2 provides a new **view on array AA**. The new view can be of different shape (as it is here). But remember that a view still points to the same data as the array on which it is view.

The rules Numpy uses when it doesn't and when it does make a copy of the data are not trivial. 

Be sure to keep this phenomenon in the back of your mind when confronted with a nasty bug in your code.

## Tricks with Arrays

### Exercise

Given the array C of shape (128,) make it into an array of shape (128,1). You can do that with the reshape method or with the use of the 'newaxis' index.

```python
u = C.copy()

u[:,np.newaxis].shape
```

Output:
```
(128, 1)
```

<!-- grade_id: cell-8125b850b25433b1, points: 0, grade: false, solution: true -->
```python
def ncomma_to_ncomma1(a):
    # YOUR CODE HERE
    return a[:,np.newaxis]
    raise NotImplementedError()
```

<!-- grade_id: cell-07b408ec9cdbbaf1, points: 1, grade: true, solution: false -->
```python
assert ncomma_to_ncomma1(C).shape == (128,1)
assert all(ncomma_to_ncomma1(C).flatten() == C)
n = randint(50,500)
D = rand(n)
assert ncomma_to_ncomma1(D).shape == (n,1)
assert all(ncomma_to_ncomma1(D).flatten() == D)
```

We make some new data to work on:

```python
seed(38293804)
A35 = randint(1,10,size=(3,5))
print(A35)

v5 = array([1,2,3,4,5])
print(v5)

v3 = array([1,2,3])
print(v3)
```

Output:
```
[[2 9 5 7 2]
 [4 3 4 2 7]
 [2 1 8 2 3]]
[1 2 3 4 5]
[1 2 3]
```

### Exercise

Subtract the (5,) array from each of the rows of A35. **Note there is no need to first duplicate the v5 array to form a (3,5) shaped array**. Note that your function should work for all arrays of size (m,n) and rows of size (n,).

```python
t = A35.copy()
q = v5.copy()
q = q[:np.newaxis,]
print(t.shape, q.shape)
t - q
```

Output:
```
(3, 5) (5,)
array([[ 1,  7,  2,  3, -3],
       [ 3,  1,  1, -2,  2],
       [ 1, -1,  5, -2, -2]])
```

<!-- grade_id: cell-8e706889deb98346, points: 0, grade: false, solution: true -->
```python
def subtract_row(a, r):
    """Subtract row r from all rows in a"""
    # YOUR CODE HERE
    return a - r[: np.newaxis,]
    raise NotImplementedError()
```

<!-- grade_id: cell-f93ee2bc1abadd8a, points: 1, grade: true, solution: false -->
```python
assert all(subtract_row(A35, v5) == array([[ 1,  7,  2,  3, -3],
       [ 3,  1,  1, -2,  2],
       [ 1, -1,  5, -2, -2]]))
```

### Exercise

Subtract the (3,) array v3 from each of the columns of A35.  Note that your function should work for all arrays of size (m,n) and columns of size (m,).

```python
t = A35.copy()
s = v3.copy()
s = s[:,np.newaxis]
print(t.shape, s.shape)
t - s
```

Output:
```
(3, 5) (3, 1)
array([[ 1,  8,  4,  6,  1],
       [ 2,  1,  2,  0,  5],
       [-1, -2,  5, -1,  0]])
```

<!-- grade_id: cell-67d1fe25f44ec2a8, points: 0, grade: false, solution: true -->
```python
def subtract_col(a, c):
    # YOUR CODE HERE
    return a - c[:, np.newaxis]
    raise NotImplementedError()
```

<!-- grade_id: cell-500e37762b09ce8d, points: 1, grade: true, solution: false -->
```python
assert all( subtract_col(A35, v3) == \
array([[ 1,  8,  4,  6,  1],
       [ 2,  1,  2,  0,  5],
       [-1, -2,  5, -1,  0]]))

```

## Linear Algebra

Python supports many linear algebra functions like calculating norm of a vector or determinant of a matrix. 

In Python 3 the `@` operator for matrix multiplication was introduced. This means that `A @ B` denotes the matrix multiplication of a matrix (array) A of shape (m,n) with a matrix (array) B of shape (n,k).

Python also allows arrays of shape (n,) to be used in matrix multiplications. Depending on the context it depends on whether Python interprets an array with shape (n,) as a matrix of shape (n,1) or (1,n).

Below some examples are given.

```python
# creation of some random matrices and vectors. Note difference in size (3,) and (3,1)

seed(324893485)
A = randint(1,10,size=(3,4))
B = randint(1,10,size=(3,3))
x = randint(1,10,size=(4,))
y = randint(1,10,size=(3,))
v = x.reshape((4, 1))
z = randint(1, 10, size=(4,1))
w = y.reshape((3, 1))
```

```python
# norm and determinant

print('norm y:', linalg.norm(y))
print('determinant B:', linalg.det(B))
```

Output:
```
norm y: 11.916375287812984
determinant B: 53.999999999999986
```

```python
# see difference in outcome

print(A @ x)
print(A @ v)
```

Output:
```
[ 88 118  60]
[[ 88]
 [118]
 [ 60]]
```

```python
# alignment problem will occur

A @ y
```

Output:
```
ValueError                                Traceback (most recent call last)
<ipython-input-315-f5d970293c6e> in <module>
      1 # alignment problem will occur
      2 
----> 3 A @ y

ValueError: matmul: Input operand 1 has a mismatch in its core dimension 0, with gufunc signature (n?,k),(k,m?)->(n?,m?) (size 3 is different from 4)
```

```python
# This should work

print(y @ A)
```

Output:
```
[ 76  92 132  85]
```

```python
# See difference in outcome. Understand difference

# v.Transpose @ z gives dot product
print(v.T@z)

# v @ z.Transpose result in matrix
print(v@z.T)
```

Output:
```
[[130]]
[[14 12 16 14]
 [21 18 24 21]
 [49 42 56 49]
 [42 36 48 42]]
```

So nothing for you to do here except note that an array of shape (n,) can be used as either row vector of column vector in a vector-matrix or matrix-vector multiplication respectively.

## End of Notebook

```python

```